<a href="https://colab.research.google.com/github/mahmud-nobe/CS110_pcw_1.2/blob/master/cs110_pcw_1_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Md Mahmudunnobe"
COLLABORATORS = "-"
```

---

# CS110 Pre-class Work 1.2





## Question 1
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](https://github.com/mahmud-nobe/CS110_pcw_1.2/raw/master/insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
Line 7 decreses the value of i by 1. Here before entering the while loop we are checking the last element of the sorted part of the list, (j-1)'th element, if it is higher then 'our current element that we are checking' (A[j] or key). 
Then we continue to check with the previous elements of the sorted part of the list until either we reached to the 1st element or the element value is not higher then the key.

So, in line 7, we are decreasing the index as we are going from the last index of the sorted part of the list downto the 1st index of that part.

If we didn't have line 7, we would check only the last element of the sorted part (which is also the adjacent element of our current 'key') every time and the while loop would not terminate at all if that element is higher than the 'key'. 

## Question 2
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
            A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertionSort(A):
    # before we are starting the for loop from 0, the first element of the list,
    # while it doesn't hamper the accuracy of the algorithm, but we can totally skip
    # the first element and start from the second element (index no 1) as the pseudo code suggested
    for j in range(1,len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        # before the code is swaping the value of key within the already sorted elements
        # inside of the while loop. But while loop works when the element value is higher than the key.
        # But we want to swap only when the while loop terminates, which means when we are sure that there is no higher element left 
        # behind this value.
        # Also when we swap it again and again inside of the while loop, the algorithm has to do extra unnecessary steps which decrease its efficiency.
        A[i+1] = key  
    return A
    raise NotImplementedError()
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertionSort([0]) == [0])
assert(insertionSort([-1,1]) == [-1,1])
assert(insertionSort([1,-1]) == [-1,1])
assert(insertionSort([1,6,3,6]) == [1,3,6,6])
```

```python
insertionSort([-1,1]) == [-1,1]
```

Output:
```
True
```

```python
insertionSort([1,6,3,6]) == [1,3,6,6]
```

Output:
```
True
```

```python
def insertionSort(A):
    for j in range(1,len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        print("j:",j, ", i after termination:", i, ", A[i]: ", A[i])
        A[i+1] = key
        print(A)
    return A
```

```python
insertionSort([1,6,3,6])
```

Output:
```
j: 1 , i after termination: 0 , A[i]:  1
[1, 6, 3, 6]
j: 2 , i after termination: 0 , A[i]:  1
[1, 3, 6, 6]
j: 3 , i after termination: 2 , A[i]:  6
[1, 3, 6, 6]
[1, 3, 6, 6]
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
I checked with some test cases and everytime it outputs the sorted list. Also the code is now matched with the abovementioned pseudocode. I also use some print statement to see how the algorithm is working.




## Question 3 
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertionSort`,  your new Python function, `bubbleSort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort(A):
    for i in range(len(A)):
      j = len(A)-1
      while (j > i):
        if (A[j] < A[j-1]):
          temp = A[j]
          A[j] = A[j-1]
          A[j-1] = temp
        j -= 1
    return A

    raise NotImplementedError()
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->

Some test cases are given below:



```python
bubbleSort([1,6,3,6])

```

Output:
```
[1, 3, 6, 6]
```

```python
bubbleSort([0, 0, -1, 8, 3])

```

Output:
```
[-1, 0, 0, 3, 8]
```

```python
bubbleSort([3,2,1,0,-1,-2])

```

Output:
```
[-2, -1, 0, 1, 2, 3]
```

This is working for these three test cases. But it's not un ambigious proof that the code is working because there could exist some extreme case for that the code might not work.

So, along with testing with some potential extreme and normal test cases, it's also a good practice to check if the loop invariant in the algorithm outputs correct result in 'initiation', 'maintenanace' and 'termination' phase.

## Question 4
The following pseudocode corresponds to a third sorting algorithm to be covered in Lesson 1.2, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free. `selectionSort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](https://github.com/mahmud-nobe/CS110_pcw_1.2/raw/master/selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selectionSort(A):
    n = len(A)
    for i in range(n-1):
      minidx = i
      for j in range(i+1, n):
        if (A[j] < A[minidx]):
          minidx = j
      temp = A[i]
      A[i] = A[minidx]
      A[minidx] = temp
    return A
    
    raise NotImplementedError()
```

```python
selectionSort([3,2,1,0,-1,-2])
```

Output:
```
[-2, -1, 0, 1, 2, 3]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**